## Before submitting
1. Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

2. Make sure that no assertions fail or exceptions occur, otherwise points will be subtracted.

3. Please submit only the `*.ipynb` file.

4. Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Edit only between `YOUR CODE HERE` and `END YOUR CODE`.

5. Fill your group name and collaborators below:

```python
GROUPNAME = ""
COLLABORATORS = "Taner Metin"
```

---

# Exercise Sheet 1: Python Basics

This first  exercise sheet tests the basic functionalities of the Python programming language in the context of a simple prediction task. We consider the problem of predicting health risk of subjects from personal data and habits. We first use for this task a decision tree

![](tree.png)

adapted from the webpage http://www.refactorthis.net/post/2013/04/10/Machine-Learning-tutorial-How-to-create-a-decision-tree-in-RapidMiner-using-the-Titanic-passenger-data-set.aspx. For this exercise sheet, you are required to use only pure Python, and to not import any module, including numpy. In exercise sheet 2, the nearest neighbor part of this exercise sheet will be revisited with numpy.

## Classifying a single instance (15 P)

* Create a function that takes as input a tuple containing values for attributes (smoker,age,diet), and computes the output of the decision tree. Should return `"less"` or `"more"`.
* Test your function on the tuple `('yes', 31, 'good')`,

<!-- grade_id: cell-b66d7278bc313c94, points: 0, grade: false, solution: true -->
```python
def decision(x):
    # >>>>> YOUR CODE HERE
    #N= tuple(raw_input().split(','))
    
    
    n0 = x[0].strip('\'')
    n1 = x[1]
    n2 = x[2].strip(' \'')

    if len(x) != 3:
        raise NotImplementedError("Enter a convenient tuple.")
    if n0 == 'yes':
        if ( n1 < 29.5) :
            return 'less'
        else:
            return 'more'
    else:
        if n2 == 'good':
            return 'less'
        else:
            return 'more'
    # <<<<< END YOUR CODE
```

<!-- grade_id: cell-c31b80471db3132f, points: 15, grade: true, solution: false -->
```python
x = ('yes', 32, 'good')
assert decision(x) == 'more'

```

## Reading a dataset from a text file (10 P)

The file `health-test.txt` contains several fictious records of personal data and habits.

* Read the file automatically using the methods introduced during the lecture.
* Represent the dataset as a list of tuples.

<!-- grade_id: cell-c1a8bc4c0e4ccb26, points: 0, grade: false, solution: true -->
```python

def gettest():
    # >>>>> YOUR CODE HERE
    with open('health-test.txt', 'r') as ftest:
        L = list()
        global count
        count = 0
        for line in ftest:
            print line
            count += 1 
            L.extend([x for x in str.split(line[:-1], ',')])
    print "number of lists:", count
    print L
    global subList
    subList = tuple([L[n:n+3] for n in range(0, len(L), 3)])
    print 'exp:', subList[5]
    print 'res:', decision(subList[5])
    
    if ftest == 0:
        raise NotImplementedError("Could not open file to read.")
    # <<<<< END YOUR CODE
```

<!-- grade_id: cell-4e1f7ad1e66b3121, points: 10, grade: true, solution: false -->
```python
gettest()
```

Output:
```
yes,21,poor

no,50,good

no,23,good

yes,45,poor

yes,51,good

no,60,good

no,15,poor

no,18,good

number of lists: 8
['yes', '21', 'poor', 'no', '50', 'good', 'no', '23', 'good', 'yes', '45', 'poor', 'yes', '51', 'good', 'no', '60', 'good', 'no', '15', 'poor', 'no', '18', 'good']
exp: ['no', '60', 'good']
res: less
```

## Applying the decision tree to the dataset (15 P)

* Apply the decision tree to all points in the dataset, and return the percentage of them that are classified as "more".

<!-- grade_id: cell-6703ef98e2b5c93b, points: 0, grade: false, solution: true -->
```python
def evaluate_testset():
    # >>>>> YOUR CODE HERE
    mores = 0
    for i in range(count):
        print decision(subList[i])
        if decision(subList[i])== 'more':
            mores = mores + 1
        if len(subList[i]) != 3:
            raise NotImplementedError("Given data is not convenient.")
     
    print "Health risk for the given dataset:%{}" .format((mores * 100)/count)
    
    
    # <<<<< END YOUR CODE
```

<!-- grade_id: cell-c13a0b23c9faba52, points: 15, grade: true, solution: false -->
```python
evaluate_testset()

```

Output:
```
more
less
less
more
more
less
more
less
Health risk for the given dataset:%50
```

## Learning from examples (10 P)

Suppose that instead of relying on a fixed decision tree, we would like to use a data-driven approach where data points are classified based on a set of training observations manually labeled by experts. Such labeled dataset is available in the file `health-train.txt`. The first three columns have the same meaning than for `health-test.txt`, and the last column corresponds to the labels.

* Write a procedure that reads this file and converts it into a list of pairs. The first element of each pair is a triplet of attributes, and the second element is the label.

<!-- grade_id: cell-fc38ed11fee6fbeb, points: 0, grade: false, solution: true -->
```python
def gettrain():
    # >>>>> YOUR CODE HERE
    with open('health-train.txt', 'r') as ftrain:
        
        T = list()
        line_counter = 0
        
        for line in ftrain:
            print line
            line_counter = line_counter +1
            T.extend([x for x in str.split(line[:-1], ',')])
            
        dataList = list()
        labelList = list()
        
        for i in range(line_counter * 4):
            if (i+1)%4 ==0:
                labelList.append(T[i])
            else:
                dataList.append(T[i])
        
        print 'dataList:{}\n'.format(dataList)
        print 'labelList:{}' .format(labelList)
        
        
        
            
        if ftrain == 0:
            raise NotImplementedError("Inconvenient type of data")
    # <<<<< END YOUR CODE
```

<!-- grade_id: cell-a3d593f232e0403a, points: 10, grade: true, solution: false -->
```python
gettrain()
```

Output:
```
yes,54,good,less

no,55,good,less

no,26,good,less

yes,40,good,more

yes,25,poor,less

no,13,poor,more

no,15,good,less

no,50,poor,more

yes,33,good,more

no,35,good,less

no,41,good,less

yes,30,poor,more

no,39,poor,more

no,20,good,less

yes,18,poor,less

yes,55,good,more

dataList:['yes', '54', 'good', 'no', '55', 'good', 'no', '26', 'good', 'yes', '40', 'good', 'yes', '25', 'poor', 'no', '13', 'poor', 'no', '15', 'good', 'no', '50', 'poor', 'yes', '33', 'good', 'no', '35', 'good', 'no', '41', 'good', 'yes', '30', 'poor', 'no', '39', 'poor', 'no', '20', 'good', 'yes', '18', 'poor', 'yes', '55', 'good']

labelList:['less', 'less', 'less', 'more', 'less', 'more', 'less', 'more', 'more', 'less', 'less', 'more', 'more', 'less', 'less', 'more']
```

## Nearest neighbor classifier (25 P)

We consider the nearest neighbor algorithm that classifies test points following the label of the nearest neighbor in the training data. For this, we need to define a distance function between data points. We define it to be

`d(a, b) = (a[0] != b[0]) + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] != b[2])`

where `a` and `b` are two tuples corrsponding to the attributes of two data points.

* Write a function that retrieves for a test point the nearest neighbor in the training set, and classifies the test point accordingly.
* Test your function on the tuple `('yes', 31, 'good')`

<!-- grade_id: cell-671ea24ec8a11241, points: 0, grade: false, solution: true -->
```python
def neighbor(x, trainset):
    # >>>>> YOUR CODE HERE
    raise NotImplementedError("Replace this line by your code.")
    # <<<<< END YOUR CODE
```

<!-- grade_id: cell-a36122337853f195, points: 15, grade: true, solution: false -->
```python
# Test
x = ('yes', 31, 'good')
assert neighbor(x, gettrain()) == "more"

```

* Apply both the decision tree and nearest neighbor classifiers on the test set, and find the data point(s) for which the two classifiers disagree, and with which probability it happens.

<!-- grade_id: cell-8dbf7da153f3d797, points: 0, grade: false, solution: true -->
```python
def compare():
    # >>>>> YOUR CODE HERE
    raise NotImplementedError("Replace this line by your code.")
    # <<<<< END YOUR CODE
    return Xdisagree, probability
```

<!-- grade_id: cell-3b55f7e89ad4dfeb, points: 10, grade: true, solution: false -->
```python
Xdisagree, probability = compare()
assert type(Xdisagree) == list
assert probability >= 0.0 and probability <= 1.0
```

One problem of simple nearest neighbors is that one needs to compare the point to predict to all data points in the training set. This can be slow for datasets of thousands of points or more. Alternatively, some classifiers train a model first, and then use it to classify the data.

## Nearest mean classifier (25 P)

We consider one such trainable model, which operates in two steps:

(1) Compute the average point for each class, (2) classify new points to be of the class whose average point is nearest to the point to predict.

For this classifier, we convert the attributes smoker and diet to real values (for smoker: yes=1.0 and no=0.0, and for diet: good=0.0 and poor=1.0), and use the modified distance function:

`d(a,b) = (a[0] - b[0]) ** 2 + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] - b[2]) ** 2`

We adopt an object-oriented approach for building this classifier.

* Implement the methods `train` and `predict` of the class `NearestMeanClassifier`.

<!-- grade_id: cell-e0b339bfd0fcc16c, points: 0, grade: false, solution: true -->
```python
class NearestMeanClassifier:
    def train(self, dataset):
        # >>>>> YOUR CODE HERE
        raise NotImplementedError("Replace this line by your code.")
        # <<<<< END YOUR CODE

    def predict(self, x):
        # >>>>> YOUR CODE HERE
        raise NotImplementedError("Replace this line by your code.")
        # <<<<< END YOUR CODE
        return prediction
```

* Build an object of class `NearestMeanClassifier`, train it on the training data, and print the mean vector for each class.

<!-- grade_id: cell-5f7f00ee83c94703, points: 0, grade: false, solution: true -->
```python
def build_and_train():
    # >>>>> YOUR CODE HERE
    raise NotImplementedError("Replace this line by your code.")
    # <<<<< END YOUR CODE
build_and_train()
```

<!-- grade_id: cell-415891bde4cbde19, points: 15, grade: true, solution: false -->
```python

```

* Predict the test data using the nearest mean classifier and print all test examples for which all three classifiers (decision tree, nearest neighbor and nearest mean) agree.

<!-- grade_id: cell-f37f3035a32a8f85, points: 0, grade: false, solution: true -->
```python
def predict_test():
    # >>>>> YOUR CODE HERE
    raise NotImplementedError("Replace this line by your code.")
    # <<<<< END YOUR CODE
    return agreed_samples
predict_test()
```

<!-- grade_id: cell-853c957eaaf81c28, points: 10, grade: true, solution: false -->
```python

```